```python
# Initialize Otter
import otter
grader = otter.Notebook("Practice.ipynb")
```

# Assignment 0: Practice - Gradescope

This assignment exists for practice only.

**Bonus:** This assignment is out of 0.25% of your Extra Credit. Anyone who completed this practice assignment will receive 0.25% EC on thier final score.

This assignment covers assignment completion, _gradescope_ submission, and anaconda downloading.


## How to complete assignments

Whenever you see:

```python
answers = ...
```

You need to replace this section with some code that answers the questions and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `...`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `grader.check("Question_id")` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. Once you run the cell, you will see the public tests we provided for you.

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

```python
# Setup for JupyterQuiz
from quiz import display_quiz, record_quiz, check_quiz_answer, show_chosen_option
from IPython.core.display import display, Javascript, HTML

import numpy as np

path = "https://raw.githubusercontent.com/COGS118A/DiscussionLabExercises/main/practice/"
```

### Q1 - Jupyter Quiz Practice Question (0.1 points)

```python
HTML(display_quiz(f"{path}A.txt"))
```

```python
show_chosen_option("A1")
```

```python
grader.check("x^3_question")
```

### Q2 - Practice Question (0.15 points)

Consider your response to the following question: which course are you MOST excited about this quarter?:
- A) COGS 118A
- B) Introduction to Supervised Machine Learning
- C) COGS1111118A
- D) Intro Sup ML

Replace `...` below to assign the letter corresponding to your response to the object `favorite_class`. Be sure that you put quotes around the the letter for your response. 

For example, if "COGS 118A" (option A above) were my answer, I would replace `favorite_class = ...` with `favorite_class = 'A'`.

After you add your answer, run the code in that cell using 'Shift + Enter' or clicking the "Run" icon above. A number will appear to the left of the cell once your code has executed.

<!-- grade_id: cell-e4f640daee64fc2e, points: 0, grade: false, solution: true -->
```python
# any answer is a good answer here!
favorite_class = ...
```

Typically, after you write some code in your assignments, there will be a few assert statements, similar to what you see below. These help to check that you're on the right path. If you are on the right track, when, you run the code in those cells, they will "pass silently", meaning you won't see any output. 

However, if you get an AssertionError, you'll want to return to the code you right and re-think your response.

_Points:_ 0.1

```python
grader.check("favorite_class")
```

**Note**: Just because an assert doesn't fail does *not* mean you've got the answer right. There are hidden tests in your assignments. These check to ensure the code you've written does what we asked. As long as you do what we ask *and* pass the assert statements, you'll receive full credit.

<!-- BEGIN QUESTION -->

## Q3 - Free Response

Describe what you want to learn in this class in a few sentences.

_Points:_ 0.1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## The End!

This is the end of the practice assignment!

Have a look back over your answers, and also make sure to `Restart & Run All` from the kernel menu to double check that everything is working properly. This restarts everything and runs your code from top to bottom.

When you are ready to submit your assignment, you can click `Validate` at the top. Note that in some assignments the code will take too long to run and validation may fail. Validation is just a final check that all the asserts are passing without failing.

Once you're happy with your work, click the disk icon to save, and submit the zip file onto gradescope. **You MUST submit all the required component to receive credit.**

Note that you can submit at any time, but **we grade your most recent submission**. This means that **if you submit an updated notebook after the submission deadline, it will be marked as late**.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Please make sure to see the output of the gradescope autograder. You are responsible for waiting and ensuring that the autograder is executing normally for your submission. Please create a campuswire post if you see errors in autograder execution. Only the images in the directory `imgs` will be zipped and submit to gradescope.

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['record.txt'])
```

Tests for favorite_class (points: 0.1)
```
>>> possible_answers = ['A', 'B', 'C', 'D']
>>> assert favorite_class in possible_answers, "Did you put in some other options?"
```


Tests for x^3_question (points: 0.05)
```
>>> # Public Test
>>> check_quiz_answer("A1", 2)
```
