```python
# Initialize Otter
import otter
grader = otter.Notebook("A4.ipynb")
```

# Assignment 4


## **Due: May 3rd (Wednesday), 2023, 8:00pm (PT)**

### **Instructions:**

Your Jupyter notebook assignment will often have 3 elements: written answers, code answers, and quiz answers. For written answers, you may insert images of your handwritten work in code cells, or write your answers in markdown and LaTeX. For quiz answers, your `record.txt` file will record your answer choices in the quiz modules for submission. Both your quiz answers and code answers will be autograded on Gradescope. This assignment does not have the quiz portion.

For all elements, DO NOT MODIFY THE CELLS. Put your answers **only** in the answer cells given, and **do not delete cells**. If you fail to follow these instructions, you will lose points on your submission.

Make sure to show the steps of your solution for every question to receive credit, not just the final answer. You may search information online but you will need to write code/find solutions to answer the questions yourself. You will submit your .ipynb file and record.txt to gradescope when you are finished.

### **Late Policy:**

Late assignments will be accepted at 75% credit up to one week late. Consult the syllabus for more info on the late policy.

### How to Include Your Math Written Answer?

You could use inline $\LaTeX$ in markdown (recommended) or use markdowns' include image functionality to submit your written responses.

#### $\LaTeX$ (recommended)
[Here is a fantastic tutorial from CalTech about using $\LaTeX$ in Jupyter Notebook.](http://chebe163.caltech.edu/2018w/handouts/intro_to_latex.html). You could also find various $\LaTeX$ tutorials and cheat sheets online.

#### Include Images
If you are still getting familiar with using LaTeX, handwrite the response on paper or the stylus. Take a picture or screenshot of your answer, and include that image in the Jupyter Notebook. Be sure to include that image in the `\imgs` directory. Let's say you have your Q1 response saved as `imgs/Q1.png`; the markdown syntax to include that image is `![Q1](imgs/Q1.png)`. 

## Important Notice

You must check both submission output on the gradescope (`Assignment 4` and `Assignment 4 - Manual Grading`) correctly reflects your work and responses. If you notice inconsistencies between your notebook and the Manual Grading portion, you need to make a campuswire post, and we can help you with that.

**Other**

If you are not feeling comfortable with the programming assignments in this homework, it might help to take a look at [https://github.com/UCSD-COGS108/Tutorials](https://github.com/UCSD-COGS108/Tutorials).

```python
%load_ext autoreload
%autoreload 2

# Run me
# !pip uninstall -y otter-grader
# !pip install -q  git+https://github.com/scott-yj-yang/otter-grader.git --no-warn-script-location
```

# Question 1: Cross Validation

For multiple choice questions, write your solution as a list of strings by replacing the "..." 

Ex.: `["A", "C"]` if you think the answers are A and C) and `["A"]` if you think the answer is A.

K-fold cross validation allows us to test and compare many combinations of hyperparameters for a given model architecture or algorithm. When performing K-fold CV for model selection and assuming we chose a reasonable value of k, which model should we choose?

A. The model that has the highest average training accuracy across all k-folds of cross validation.

B. The model with the smallest differences between the average training accuracy and average test accuracy within a given fold.

C. The model with the highest average test accuracy across folds.

D. The model with the highest test accuracy in a given fold.

_Points:_ 0.15

```python
q1_1 = ...
```

```python
grader.check("MCQ_1_1")
```

Why do we use nested cross validation? (Choose all that apply)

A. To avoid over-fitting the model on the training set.

B. To minimize the complexity of training on large data sets.

C. In order to optimize the hyper-parameters of a model.

D. To take advantage of small test sets in validating accuracy.

_Points:_ 0.15

```python
q1_2 = ...
```

```python
grader.check("MCQ_1_2")
```

What is a drawback with using nested cross validation?

A. It can lead to over-fitting of the model.

B. It has a very expensive computation relative to other methods.

C. It can lead to over-fitting of the hyper-parameters.

D. There are no drawbacks to using nested cross validation.

_Points:_ 0.15

```python
q1_3 = ...
```

```python
grader.check("MCQ_1_3")
```

<!-- BEGIN QUESTION -->

If we have a small 1000 point dataset and we want to use k-fold cross validation, what are the advantages and disadvantages of using k=4 and k=1000 for model selection?

_Points:_ 0.15

_Type your answer here, replacing this text._

<!-- END QUESTION -->


# Question 2: Error Metrics

Now consider that you want to develop a classification model to identify whether a tumor is a malignant tumor or a benign tumor. This model will assist doctors in deciding whether or not a patient needs chemotherapy. After generating predictions with your models (using the same test set), you generate two confusion matrices: A and B. Calculate the error metrics for both confusion matrices and use the metrics to justify which model is more appropriate for this task.

Confusion Matrix A

|              |          | predicted | diagnosis |       |
|--------------|----------|-----------|-----------|-------|
|              |          | postive   | negative  | total |
| true classes | positive | 15        | 42        | 57    |
| true classes | negative | 2         | 41        | 43    |
|              | total    | 17        | 83        | 100   |

Confusion Matrix B

|              |          | predicted | diagnosis |       |
|--------------|----------|-----------|-----------|-------|
|              |          | postive   | negative  | total |
| true classes | positive | 40        | 17        | 57    |
| true classes | negative | 31        | 12        | 43    |
|              | total    | 71        | 29        | 100   |

<!-- BEGIN QUESTION -->

## Question 2.1
Compute the Recall for each confusion matrix and discuss its implications in this context:

$$\text{Recall} = \frac{\text{TP}}{\text{TP + FN}}$$


_Points:_ 0.15

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 2.2

Compute the Precision for each confusion matrix and discuss its implications in this context:

$$ \text{Precision} = \frac{\text{TP}}{\text{TP+FP}} $$

_Points:_ 0.15

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 2.3

Compute the F-value for each confusion matrix, what does this mean?:

$$ \text{F-value} = \frac{\text{2 * Precision * Recall}}{\text{Precision + Recall}} $$

_Points:_ 0.15

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 2.4

With the metrics you've gathered, which model is more appropriate for classifying cancer cells, the model that generated Confusion Matrix A or B? Why?

_Points:_ 0.15

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Question 3 Model Selection

Read the following statements and identify whether they are true or false, Replace the "..." with your answer 

Ex.: `q3_5=True` if you think the answer to the fifth question is True. 

Nested K-Fold Cross Validation can be used for simultaneous model selection and validation.

_Points:_ 0.15

```python
q3_1 = ...
```

```python
grader.check("TF_3_1")
```

Using just point estimates of their mean validation performances, you can determine whether model $A$ is significantly better than model $B$

_Points:_ 0.15

```python
q3_2 = ...
```

```python
grader.check("TF_3_2")
```

Cross validation allows us to approximate a distribution over validation performance.

_Points:_ 0.15

```python
q3_3 = ...
```

```python
grader.check("TF_3_3")
```

It is impossible to over-fit a model using cross validation and grid-search.

_Points:_ 0.15

```python
q3_4 = ...
```

```python
grader.check("TF_3_4")
```

# Question 4 Decision Tree Modeling Decisions

Consider the decision tree in the below figures. Suppose your colleague trained this model on a binary classification dataset of 'x's and 'o's but it isn't performing well in production. Because you took COGS 118A, they naturally turn to you for help. Your colleague tells you that in addition to performing poorly on real world data, they expected most 'x's to be cluster in the center of the feature space.

![overfit_tree_decision_boundary](imgs/overfit_tree_decision_boundary.png)

![overfit_tree](imgs/overfit_tree.png)

Based on your knowledge about decision tree modeling and regularization choices, how might you retain this model? (Select all that apply)

A. Train with a maximum tree depth

B. Train with a minimum tree depth

C. Add an L1 regularization term

D. Add a tree pruning technique

E. Specify the minimum number of samples that can be split for any node

F. Specify the maximum number of samples that can be split for any node

G. Train with cost complexity: use a larger terminal node penalty

H. Train with cost complexity: use a smaller terminal node penalty

(Write your solution as a list of strings by replacing the "..." Ex.: `["A", "C"]` if you think the answers are A and C.)

```python
q4 = ...
```

```python
grader.check("MCQ_4")
```

<!-- BEGIN QUESTION -->

# Question 5 Training and Test metrics

Consider the following hyper-parameter search. Here $score$ (y-axis) is $R^2$ and $degree$ is some hyper-parameter. What parameter setting do you expect to generalize best to new data? Briefly explain your reasoning.

![train_test_metrics](imgs/train_test_metrics.png)

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->


# Question 6 K-Nearest Neighbors Classification

Although k-Nearest Neighbors models are the "dumbest" machine learning classifiers, they can be highly accurate and valuable in the real world. With k-NN models, the model is the data itself; for a new data point you want to predict, you compute the similarity between that point and every point in your dataset. 

The predicted class for a test set example is the majority of k classes.

![knn](imgs/knn.png)

As you can imagine, with a large dataset, the model will perform exceptionally well. The downside is that the larger the dataset gets, the k-NN model will become more and more computationally complex: $O(1)$ training time (adding 1 data point to the model) and $O(n)$ testing time where $n$ equals the size of the dataset.


### Euclidean Distances

With this computational complexity in mind, we want to make our algorithm as efficient as possible. Write a `euc_distance` function that computes the euclidean distance between a point and all other points in our training data using numpy vectorization.

_Points:_ 0.2

```python
import numpy as np
```

```python
def euc_distance(point, all_points):
    """
    Input
    -----
    point: a numpy array of shape (1,2)
    
    all_points: a numpy array of shape (n,2)
    
    Returns
    -----
    distances: a numpy array of shape (n,)
    """
    
    distances = ...
    
    return distances
```

### KNN Algorithm Implementation

Below, you are given training and testing data that is visualized in matplotlib. Your job is to implement the KNN Algorithm using your `euc_distance` function so that we can predict the class for the black and green clusters of testing data, using our blue and red training data.

```python
import numpy as np
np.random.seed(42)
mean1, cov1 = np.array([10,2]), np.array([[3,0],[5,3]]) 
mean2, cov2 = np.array([5,4]), np.array([[28,3],[3,20]]) 
cluster_1 = np.random.multivariate_normal(mean1, cov2, 100) 
cluster_2 = np.random.multivariate_normal(mean1, cov2, 100) + 10
test_data_1 = np.random.multivariate_normal(mean1, cov1, 20) 
test_data_2 = np.random.multivariate_normal(mean2, cov2, 20) + 10
test_data = np.random.multivariate_normal((mean1+mean2)/2, (cov1+cov2)/2, 100) + 5

```

```python
import matplotlib.pyplot as plt

plt.figure(figsize=(10,8), dpi=150)
plt.scatter(cluster_1[:,0], cluster_1[:,1], label="Positive Class")
plt.scatter(cluster_2[:,0], cluster_2[:,1], label="Negative Class", c='r')
plt.scatter(test_data_1[:,0], test_data_1[:,1], label="Positive Test", c='black')
plt.scatter(test_data_2[:,0], test_data_2[:,1], label="Negative Test", c='green')
plt.xlabel("X1")
plt.ylabel("X2")
plt.title("k-NN Training and Testing Data")
plt.legend()
plt.show()
```

#### Implementation Instructions / Psuedocode:

- Create a k-NN model that can take `k` number of neighbors to select for majority voting, `x_train` data points that comprise the model (the red and blue points), and `y_train` class labels corresponding to `x_train`, as an argument
- Store the distance between each `x_test` data point and the training set into an array and sort it according to ascending order of their distances
    - _Hint:_ `np.argsort` is useful
- Select the first K elements
- Perform majority voting so that the class with most occurences in the first K elements will be assigned as the new class for that test data point.
    - _Hint:_ `mode` from scipy.stats may be useful
- `return` a NumPy array of the labels generated with `x_test`


_Points:_ 0.5

```python
from scipy.stats import mode


X_train = np.vstack((cluster_1, cluster_2))
y_train = np.append(np.ones(len(cluster_1)), np.zeros(len(cluster_2))).reshape(len(cluster_1) + len(cluster_2), 1)
x_test = np.vstack((test_data_1, test_data_2))
y_test = np.append(np.ones(len(test_data_1)), np.zeros(len(test_data_2))).reshape(len(test_data_1) + len(test_data_2), 1)

def predict(k, x_train, y_train, x_test):
    """
    Input
    -----
    k: number of nearest neighbors
    x_train: training data of x1,x2 coordinate positions. (n,2) numpy array 
    y_train: training data of class labels. (n,1) numpy array
    y_test: testing data of x1, x2 coordinate positions. (n,2) numpy array
    
    
    Output
    -----
    y_pred: predictions of which class all points in y_test belong to. (n,) numpy array
    
    """
    y_pred = []
    
    # Enter your implementation below
    y_pred = np.array(y_pred)
    return y_pred


def compute_accuracy(y_pred, y_test):
    incorrect = np.count_nonzero(y_pred.flatten() != y_test.flatten())
    accuracy = (1 - incorrect/len(y_test))
    return accuracy
    

y_pred_k1 = predict(1, X_train, y_train,x_test)
y_pred_k3 = predict(3, X_train, y_train,x_test)
y_pred_k5 = predict(5, X_train, y_train,x_test)


for k, predictions in zip([1,3,5],[y_pred_k1, y_pred_k3, y_pred_k5]):
    accuracy = compute_accuracy(predictions, y_test) 
    print(f"K = {k}: Accuracy = {accuracy}")

```

```python
import scipy
scipy.__version__
```

# End of A4

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Please make sure to see the output of the gradescope autograder. You are responsible for waiting and ensuring that the autograder is executing normally for your submission. Please create a campuswire post if you see errors in autograder execution.

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['imgs'])
```

Tests for KNN (points: 0.5)


Tests for MCQ_1_1 (points: 0.15)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q1_1, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for MCQ_1_2 (points: 0.15)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q1_2, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for MCQ_1_3 (points: 0.15)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q1_3, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for MCQ_4 (points: 0.3)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D", "E", "F", "G", "H"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q4, all_options), "Is your answer within the option of A-H?"
```


Tests for TF_3_1 (points: 0.15)
```
>>> # sanity check
>>> assert isinstance(q3_1, bool), "your answer should be a boolean value (True or False)"
```


Tests for TF_3_2 (points: 0.15)
```
>>> # sanity check
>>> assert isinstance(q3_2, bool), "your answer should be a boolean value (True or False)"
```


Tests for TF_3_3 (points: 0.15)
```
>>> # sanity check
>>> assert isinstance(q3_3, bool), "your answer should be a boolean value (True or False)"
```


Tests for TF_3_4 (points: 0.15)
```
>>> # sanity check
>>> assert isinstance(q3_4, bool), "your answer should be a boolean value (True or False)"
```


Tests for euclidean distance (points: 0.2)
