# Assignment 1


## **Due: April 12th (Wednesday), 2023, 8pm (Pacific Time)**

### **Instructions:**

Your Jupyter notebook assignment will often have 3 elements: written answers, code answers, and quiz answers. For written answers, you may insert images of your handwritten work in code cells, or write your answers in markdown and LaTeX. For quiz answers, your `record.txt` file will record your answer choices in the quiz modules for submission. Both your quiz answers and code answers will be autograded on Gradescope. This assignment does not have the quiz portion.

For all elements, DO NOT MODIFY THE CELLS. Put your answers **only** in the answer cells given, and **do not delete cells**. If you fail to follow these instructions, you will lose points on your submission.

Make sure to show the steps of your solution for every question to receive credit, not just the final answer. You may search information online but you will need to write code/find solutions to answer the questions yourself. You will submit your .ipynb file and record.txt to gradescope when you are finished.

### **Late Policy:**

Late assignments will be accepted at 75% credit up to one week late. Consult the syllabus for more info on the late policy.

### How to Include Your Math Written Answer?

You could use inline $\LaTeX$ in markdown (recommended) or use markdowns' include image functionality to submit your written responses.

#### $\LaTeX$ (recommended)
[Here is a fantastic tutorial from CalTech about using $\LaTeX$ in Jupyter Notebook.](http://chebe163.caltech.edu/2018w/handouts/intro_to_latex.html). You could also find various $\LaTeX$ tutorials and cheat sheets online.

#### Include Images
If you are still getting familiar with using LaTeX, handwrite the response on paper or the stylus. Take a picture or screenshot of your answer, and include that image in the Jupyter Notebook. Be sure to include that image in the `\imgs` directory. Let's say you have your Q1 response saved as `imgs/Q1.png`; the markdown syntax to include that image is `![Q1](imgs/Q1.png)`. 

## Important Notice

Only the response in the cell of _`Type your answer here, replacing this text.`_ will be recorded. Any additional cell will not be graded.

You must check both submission output on the gradescope (`Assignment 1 - Notebook` and `Assignment 1 - Manual Grading`) correctly reflects your work and responses. If you notice inconsistencies between your notebook and the Manual Grading portion, you need to make a campuswire post, and we can help you with that.

**Other**

If you are not feeling comfortable with the programming assignments in this homework, it might help to take a look at [https://github.com/UCSD-COGS108/Tutorials](https://github.com/UCSD-COGS108/Tutorials).

```python
%load_ext autoreload
%autoreload 2

# Run me
!pip uninstall -y otter-grader
!pip install -q  git+https://github.com/scott-yj-yang/otter-grader.git --no-warn-script-location
    
# Initialize Otter
import otter
grader = otter.Notebook("A1.ipynb")
```

<!-- BEGIN QUESTION -->

## Question 1: Basic Calculus

### 1.1 Derivatives with Scalars
1. $f(x) = 2\lambda + e^{\lambda x}$, where $\lambda$ is a constant, derive $\frac{\partial f(x)}{\partial x}$.

2. $f(x) = \ln (3 - e^x)$, derive  $\frac{\partial f(x)}{\partial x}$

**Note:** Only the response in the cell of _`Type your answer here, replacing this text.`_ will be recorded. Any additional cell will not be graded.

_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->


### 1.2 Derivatives with Vectors
Several particular vector derivatives are useful for the course. For matrix **A**, column vector **x** and **a**, we have:

- $\frac{\partial \mathbf a^T \mathbf x}{\partial \mathbf x} = \frac{\partial \mathbf x^T \mathbf a}{\partial \mathbf x} = \mathbf a$ 
- $\frac{\partial \mathbf x^T \mathbf A \mathbf x}{\partial \mathbf x} = (\mathbf A + \mathbf A^T) \mathbf x$. If $\mathbf A$ is symmetric, $\frac{\partial \mathbf x^T \mathbf A \mathbf x}{\partial \mathbf x} = 2 \mathbf A \mathbf x$.


The above rules adopt a denominator-layout notation. For more rules, you can refer to [this Wikipedia page](https://en.wikipedia.org/wiki/Matrix_calculus#Scalar-by-vector_identities). Please apply the above rules and calculate the following derivatives:

1. $f(x) = \lambda (1- \mathbf a^T \mathbf x)$, where $\lambda$ is a constant, derive $\frac{\partial f(\mathbf x)}{\partial \mathbf x}$.

2. $f(x) = \lambda (1 - \mathbf x^T \mathbf A \mathbf x)$ where $A$ is a symmetric matrix and $\lambda$ is a constant, derive  $\frac{\partial f(\mathbf x)}{\partial \mathbf x}$.


**Note:** Only the response in the cell of _`Type your answer here, replacing this text.`_ will be recorded. Any additional cell will not be graded.

_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 2: Solving Equations

1. Given a system of linear equations:

    $$
    \begin{cases}
    -3 w_1 + 4 w_2 & = 1 \\
    w_1 - 2 w_2 & = -1
    \end{cases}
    $$

     a) Solve for $(w_1, w_2)$.

     b) Once you obtain $(w_1, w_2)$ in Part a, please calculate final result $y$ from $y = w_1x_1 + w_2x_2$ when $(x_1, x_2) = (2, 3)$.

2. Given a system of equations:

 $$
    \begin{cases}
    -3 w_1 + 4 w_2 & = b \\
    w_1 - 2 w_2 & = -b \\
    w_1^2 + w_2^2 & = 1
    \end{cases}
    $$

Solve for $(w_1, w_2, b)$.

**Note:** Only the response in the cell of _`Type your answer here, replacing this text.`_ will be recorded. Any additional cell will not be graded.

_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 3: One-Hot Encoding

A dataset $S$ is denoted as $S=\{{\bf x}_1,{\bf x}_2,{\bf x}_3,{\bf x}_4,{\bf x}_5 \}$, where each sample refers to the specification of a car.

Represent this dataset $S$ using a matrix and show how your one-hot encoding is derived. 

**Hint**: (1) For the categorical features, you may use the one-hot encoding strategy. (2) You may choose either a row vector or a column vector to represent each data sample in your result. If you use a row vector to represent each data sample, the shape of the resulting matrix should be $5\times9$.

|          | Length (inch) | Height (inch) | Make   | Color  |
|----------|----------------|----------------|--------|--------|
| ${\bf x}_1$ | 183            | 62             | Toyota | Blue   |
| ${\bf x}_2$ | 181            | 65             | BMW    | Silver |
| ${\bf x}_3$ | 182            | 59             | BMW    | Red    |
| ${\bf x}_4$ | 179            | 68             | Ford   | Blue   |
| ${\bf x}_5$ | 182            | 53             | Toyota | Black  |


**Note:** Only the response in the cell of _`Type your answer here, replacing this text.`_ will be recorded. Any additional cell will not be graded.

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Question 4: Multiple Choice Question
Select the correct option(s). Note that there might be multiple correct options.

1. For two monotonically increasing functions $f(x)$ and $g(x)$:

    A. $f(x)+g(x)$ is always monotonically increasing.  
    B. $f(x)-g(x)$ is always monotonically increasing.  
    C. $f(x^2)$ is always monotonically increasing.  
    D. $f(x^3)$ is always monotonically increasing.  


2. For a function $f(x) = x(10-x), x\in\mathbb{R}$, please choose the correct statement(s) below:

    A. $\arg\max_{x} f(x) = 5$.  
    B. $\arg\min_{x} f(x) = 25$.  
    C. $\min_{x} f(x) = 5$.  
    D. $\max_{x} f(x) = 25$.  


3. Assume we have a function $f(x)$ which is differentiable at every $x\in \mathbb{R}$. There are three properties that describe the function $f(x)$:
    (1) $f(x)$ is a convex function.
    (2) When $x = x_{0}$, $f'(x_{0})=0$.
    (3) $f(x_{0})$ is a global minimum of $f(x)$.
    
    Which one of the following statements is **wrong**?  
    **Hint:** You can use a failure case to disprove a statement.
    
    A. Given (1) and (2), we can prove that (3) holds.  
    B. Given (2) and (3), we can prove that (1) holds.  
    C. Given (1) and (3), we can prove that (2) holds.  

**Note:** The `grader.check()` only checks whether you have responded reasonably (Only the `A/B/C/D` answer). This check will not guarantee the correctness of your answer.

_Points:_ 0.3

```python
q4_1 = []
q4_1 = ...
```

```python
q4_2 = []
q4_2 = ...
```

```python
q4_3 = []
q4_3 = ...
```

```python
grader.check("conceptual")
```

<!-- BEGIN QUESTION -->

## Question 5: Convexity

Identify the convexity for the following six functions. Which functions are **convex**? Which are **non-convex**? Which are **strictly convex**? Simply write the letters for each function graph in your answer.

![convexity](imgs/convexity.png)

**Note:** Only the response in the cell of _`Type your answer here, replacing this text.`_ will be recorded. Any additional cell will not be graded.

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 6: Basic Plots Using Matplotlib

Matplotlib is a very useful library to plot graphs. Later in the course, you may need to plot your own graphs in your report such as: loss vs. iteration, accuracy vs. category. We will start with a simple exercise.

This code creates a graph with the original sine curve and another curve where 5 iterations of Gaussian noise (0 mean, 0.1 standard deviation) are added to the original sine curve.

```python
import numpy as np
import matplotlib.pyplot as plt

np.random.seed(0)
space = np.linspace(0, 10, num = 50)
sine = np.sin(space)
sine_5 = sine

for i in range(5):
    sine_5 = sine_5 + np.random.normal(scale = 0.1, size = 50)
    
plt.scatter(space, sine, color = 'b', label = 'sine curve')
plt.plot(space, sine_5, color = 'r', label = 'noise_5_iters')
plt.legend(loc = 'upper right')
plt.show()
```

In the cell below, copy the approach above but add two new curves to the plot. Each should have a different color to distinguish it from the other curves. The new curves will add 20 and 100 iterations of Gaussian noise (0 mean, 0.1 standard deviation) onto the original sine curve. Plot the two new curves you obtained as well as the curve from the original code. The legend should label the new curves as noise 20 iters and noise 100 iters. Report your code and the final plot: one figure with four differently colored curves (original, 5 iters, 20 iters, 100 iters).

_Points:_ 0.4

```python
import numpy as np
import matplotlib.pyplot as plt
np.random.seed(0)
space = np.linspace(0, 10, num = 50)
sine = np.sin(space)

...
```

<!-- END QUESTION -->

## Question 7: Data Manipulation

```python
import matplotlib.pyplot as plt
from sklearn import datasets

# Load Iris Dataset
iris = datasets.load_iris()
X = iris.data
Y = iris.target
```

Some useful instructions are shown below:

- Get a row or a column of the array X

```python
print(X[0]) # Print the first row of array X.
print(X[:, 0]) # Print the first column of array X.
# ':' here means all rows and '0' means column 0.
```

- Get part of the array

```python
print(X[3:5, 1:3]) # Print 4th and 5th rows, 2nd and 3rd columns.
print(X[:3, :2]) # Print first 3 rows, first 2 columns.
```

### Question 7.1 
Show the first 2 features of the first 5 data points (i.e. first 2 columns and first 5 rows) of array X. (You can print the 5×2 array).

_Points:_ 0.1

```python
def q1(array):
    array = ...
    return array
```

### Question 7.2
Calculate the mean and the variance of the 2nd feature (the 2nd column) of array X.

_Points:_ 0.1

```python
def q2(array):
    ...
    return mean, var
```

### Question 7.3
 Perform a linear projection on the 1st 3 features of all data points using the weight vector $\mathbf w$, where $\mathbf w = (3, 2, 1)$. You can do so by calculating a dot product between the 1st 3 features of all data points and the weight vector $\mathbf w$. The shape of projected data points should be (150, 1) or (150,), depending on the weight vector is regarded as a matrix or a vector. Calculate the mean of the projected data points.
 
 Hint: You can calculate the projection with a single line code using np.dot, but you should be careful with the dimension matching for the dot product.

_Points:_ 0.1

```python
w = np.array([3,2,1])
```

```python
def q3(array, projection_vector):
    ...
    return mean
```

 ### Question 7.4
 Randomly sample 3 data points (rows) of array X by randomly choosing the row indices. Show the indices and the sampled data points.

Hint: np.random.randint

_Points:_ 0.1

```python
def q4(array):
    ...
    return matrix
```

### Question 7.5
Add one more feature (one more column) to the array X after the last feature. The values of the added feature for all data points are constant 1. Show the first data point (first row) of the new array.

Hint: np.ones, np.hstack

_Points:_ 0.1

```python
def q5(array):
    ...
    return row
```

### Question 7.6
Add one more data point (one more row) to the array X after the last data point. The value of the added data point is the same as the first data point. Show the first feature (first column) of the new array.

Hint: `np.vstack`

_Points:_ 0.1

```python
def q6(array):
    ...
    return matrix
```

# End of A1

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Please make sure to see the output of the gradescope autograder. You are responsible for waiting and ensuring that the autograder is executing normally for your submission. Please create a campuswire post if you see errors in autograder execution. Only the images in the directory `imgs` will be zipped and submit to gradescope.

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['imgs'])
```

Tests for conceptual (points: 0.3)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q4_1, all_options), "Is your answer within the option of A/B/C/D?"
>>> assert check_valid(q4_2, all_options), "Is your answer within the option of A/B/C/D?"
>>> assert check_valid(q4_3, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for data manipulation1 (points: 0.1)


Tests for data manipulation2 (points: 0.1)


Tests for data manipulation3 (points: 0.1)


Tests for data manipulation4 (points: 0.1)


Tests for data manipulation5 (points: 0.1)


Tests for data manipulation6 (points: 0.1)
