```python
# Initialize Otter
import otter
grader = otter.Notebook("A3.ipynb")
```

# Assignment 3


## **Due: April 26th (Wednesday), 2023, 8:00pm (PT)**

### **Instructions:**

Your Jupyter notebook assignment will often have 3 elements: written answers, code answers, and quiz answers. For written answers, you may insert images of your handwritten work in code cells, or write your answers in markdown and LaTeX. For quiz answers, your `record.txt` file will record your answer choices in the quiz modules for submission. Both your quiz answers and code answers will be autograded on Gradescope. This assignment does not have the quiz portion.

For all elements, DO NOT MODIFY THE CELLS. Put your answers **only** in the answer cells given, and **do not delete cells**. If you fail to follow these instructions, you will lose points on your submission.

Make sure to show the steps of your solution for every question to receive credit, not just the final answer. You may search information online but you will need to write code/find solutions to answer the questions yourself. You will submit your .ipynb file and record.txt to gradescope when you are finished.

### **Late Policy:**

Late assignments will be accepted at 75% credit up to one week late. Consult the syllabus for more info on the late policy.

### How to Include Your Math Written Answer?

You could use inline $\LaTeX$ in markdown (recommended) or use markdowns' include image functionality to submit your written responses.

#### $\LaTeX$ (recommended)
[Here is a fantastic tutorial from CalTech about using $\LaTeX$ in Jupyter Notebook.](http://chebe163.caltech.edu/2018w/handouts/intro_to_latex.html). You could also find various $\LaTeX$ tutorials and cheat sheets online.

#### Include Images
If you are still getting familiar with using LaTeX, handwrite the response on paper or the stylus. Take a picture or screenshot of your answer, and include that image in the Jupyter Notebook. Be sure to include that image in the `\imgs` directory. Let's say you have your Q1 response saved as `imgs/Q1.png`; the markdown syntax to include that image is `![Q1](imgs/Q1.png)`. 

## Important Notice

You must check both submission output on the gradescope (`Assignment 3` and `Assignment 3 - Manual Grading`) correctly reflects your work and responses. If you notice inconsistencies between your notebook and the Manual Grading portion, you need to make a campuswire post, and we can help you with that.

**Other**

If you are not feeling comfortable with the programming assignments in this homework, it might help to take a look at [https://github.com/UCSD-COGS108/Tutorials](https://github.com/UCSD-COGS108/Tutorials).

```python
%load_ext autoreload
%autoreload 2

# Run me
# !pip uninstall -y otter-grader
# !pip install -q  git+https://github.com/scott-yj-yang/otter-grader.git --no-warn-script-location
```

# Question 1: Conceptual Questions

Multiple Choice Section

Write your solution as a list of strings by replacing the "..." 

Ex.: `["A", "C"]` if you think the answers are A and C). 

## 1. Linear regression

Choose <b>all</b> the valid answers to the description about <b>linear regression</b> and <b>logistic regression</b> from the options below:

A. A linear regression model is linear (i.e., the features are combined linearly), while a logistic regression model is non-linear because its decision boundary is not a straight line.

B. Both linear regression and logistic regression involve minimizing some 'loss'.

C. Linear regression can only be solved by using closed-form solution while logistic regression does not have a closed-form solution.

D. The output of a linear regression predictor can be any real value, while a logistic regression predictor can only give binary outputs (i.e. either -1 or +1) and it can never estimate the probability of an instance being a positive or negative case. 

_Points:_ 0.2

```python
q1_1 = ...
```

```python
grader.check("MCQ_LR_1")
```

## 2. Gradient Descent

Choose <b>all</b> the valid answers to the description about <b>gradient descent</b> from the options below:

A. The general gradient descent method does not guarantee that the global minimum can always be reached, but there exists an algorithm that can always find the suitable learning-rate/step-size so that the global minimum can be reached after finite number of iterations.

B. Loss is always decreased as long as we are moving in the opposite direction of its gradient at each step, while the step-size does not matter.

C. With same initial weights and same training dataset, the algorithm should either reach the global minumum or end up at the same local minimum; the order in which data is fed does not matter.

D. When the learning rate is high, though training may be faster, learning may 'jump' back and forth over the minima, thus make it even more difficult to find the optimal solution.

_Points:_ 0.2

```python
q1_2 = ...
```

```python
grader.check("MCQ_LR_2")
```

## 3. Classification

Choose the <b>most significant</b> difference between regression and classification:

A. Classification can handle categorical variables directly but regression can’t.

B. Classification is supervised machine learning algorithm but regression is unsupervised machine learning algorithm.

C. Classification's output is bounded (between -1 and 1) while regression can output value less than -1 or greater than 1.

D. Classification aims to output discrete values (or categories) but regression aims to output continuous values (or numerical values).

_Points:_ 0.2

```python
q1_3 = ...
```

```python
grader.check("MCQ_LR_3")
```

## 4. Classification (example)

Assume we have a binary classification model:

\begin{align*}
f(\mathbf{x})&=
\begin{cases}
 1, & w \cdot x + b \ge 0, \\
-1, & w \cdot x + b < 0. 
\end{cases}
\end{align*}

where there’s a feature vector $x = (x_1, x_2) \in R^2$, bias $b \in R$, and a weight vector $w = (w_1, w_2) \in R^2$. The decision boundary of the classification model is: 

\begin{align*}
 w \cdot x + b = 0
\end{align*}

Part 1)  If the predictions of the classifier f and its decision boundary $w \cdot x + b = 0$ are shown in Figure 1, which one below can be a possible solution of weight vector $w$ and bias $b$? 

A. $w = (+1, +1); b = 0$

B. $w = (+1, -1); b = 0$

C. $w = (-1, -1); b = 0$

D. $w = (-1, +1); b = 0$

![Q1_reg_1](imgs/Q1_reg_1.png)

_Points:_ 0.2

```python
q1_4_1 = ...
```

```python
grader.check("MCQ_LR_4_1")
```

Part 2)  If the predictions of the classifier f and its decision boundary $w \cdot x + b = 0$ are shown in Figure 2, which one below can be a possible solution of weight vector $w$ and bias $b$? 

A. $w = (-1, +1); b = +1$

B. $w = (-1, +1); b = -1$

C. $w = (+1, -1); b = +1$

D. $w = (+1, -1); b = -1$

![Q1_reg_2](imgs/Q1_reg_2.png)

_Points:_ 0.2

```python
q1_4_2 = ...
```

```python
grader.check("MCQ_LR_4_2")
```


# Question 2: Logistic Regression Inference

We have a logistic regression classifier for a 2-dimensional feature vector $\mathbf{x}=(x_1,x_2)$:
\begin{align*}
f(\mathbf{x})&=
\begin{cases}
 1, & \frac{1}{1+e^{-(3x_1-4x_2+5)}} \geq 0.5, \\
-1, & \text{otherwise}. 
\end{cases}
\end{align*}

where $f(\mathbf{x})\in \{-1,1\}$ is the predicted label. Please solve the following problems. 

<!-- BEGIN QUESTION -->

## Question 2.1
Draw the decision boundary of the logistic regression classifier and shade the region where the classifier predicts 1. Make sure you have marked the $x_1$ and $x_2$ axes and the intercepts on those axes.

_Points:_ 0.3

```python
import numpy as np
import matplotlib.pyplot as plt

# the x1 and x2 of points on the decision boundary
boundary_x1s = np.linspace(-2, 2, 20)
# STEP 1
boundary_x2s = ...

# plot the boundary
plt.plot(boundary_x1s, boundary_x2s)
plt.xlim([-2.2, 2.2])
plt.ylim([-2.2, 2.2])
plt.annotate('x1', [2, 0.2], weight='bold')
plt.annotate('x2', [0.2, 2], weight='bold')

# generate some random data points
np.random.seed(0)
N_samples = 60
sample_x1s = np.random.rand(N_samples) * 4 - 2 
sample_x2s = np.random.rand(N_samples) * 4 - 2

# predict the labels (1 or -1) of the samples
# STEP 2
logit = ...
# STEP 3
probabilities = ...
positive = probabilities >= 0.5
negative = probabilities < 0.5

# plot the samples
plt.scatter(sample_x1s[positive], sample_x2s[positive], color="g", marker="+", label='positive')
plt.scatter(sample_x1s[negative], sample_x2s[negative], color="r", marker="_", label='negative')
plt.legend(bbox_to_anchor=(1.05, 1.0), loc='upper left')

ax = plt.gca()
ax.spines['top'].set_color('none')
ax.spines['bottom'].set_position('zero')
ax.spines['left'].set_position('zero')
ax.spines['right'].set_color('none')
ax.set_xticks( [-2, -1, 0, 1, 2])
ax.set_yticks( [-2, -1, 1, 2])
ax.set_aspect(1)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 2.2

Solve this problem using handwritten / LaTeX math.
For the new data point $\mathbf{x}_1=(3,4)$:

(1) Predict its label;

(2) Give the estimated probability of it being a <b>negative</b> sample (round your answer to 0.01);

(3) Compute its distance to the decision boundary.

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Question 3. Logistic Regression

Assume in a binary classification problem, we need to predict a binary label $y\in \{-1,1\}$ for a feature vector $x=[x_0,x_1]^\top$. In logistic regression, we can reformulate the binary classification problem in a probabilistic framework: We aim to model the distribution of classes given the input feature vector $x$. Specifically, we can express the conditional probability $p(y|x)$ parameterized by $(\mathbf w, b)$ using logistic function as:
\begin{align*}
\label{eq:logistic-regression-model}
p(y|x)=\frac{1}{1+e^{-y(\mathbf w \cdot x + b)}}
\end{align*}

where $\mathbf{w}=[w_0,w_1]^\top$ is the weight vector, and $b$ is the bias scalar. Given a training dataset $S_\text{training} = \{(x_i, y_i)\}, i=1,\ldots,n\}$, we wish to optimize the negative log-likelihood loss $\mathcal{L}(\mathbf w, b)$:

$$ \mathcal{L}(\mathbf w, b)=-\sum_{i=1}^{n} \ln p(y_i|x_i) $$

and find the optimal weight vector $\mathbf{w}$ and bias $b$ to build the logistic regression model:
$$ \mathbf{w}^*, b^* = \arg\min_{\mathbf{w}, b}\mathcal{L}(\mathbf{w}, b) $$

<!-- BEGIN QUESTION -->

## Question 3.1
In this problem, we attempt to obtain the optimal parameters $\mathbf{w}^*$ and $b^*$ by using a standard gradient descent algorithm. Assume $p_i = p(y_i|x_i)$, solve for the gradient for $\mathbf w$ and the gradient for b.

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 3.2
In reality, we typically tackle this problem in a matrix form: First, we represent data points as matrices $X=[\mathbf{x}_1, \mathbf{x}_2, \ldots, \mathbf{x}_n]^T$ and $Y=[y_1, y_2, \ldots, y_n]^T$. Thus, the negative log-likelihood loss $\mathcal{L}(\mathbf{w},b)$ can be formulated as:

\begin{align*}
\mathbf{p} = \text{sigmoid}(Y \circ (X\mathbf{w} + b\mathbf{1})),  \quad\quad \mathcal{L}(\mathbf w, b)=-\mathbf{1}^T \ln \mathbf{p}
\end{align*}

where $\mathbf{1} = [1,1,...,1]^T \in \mathbb{R}^n$ is a $n$-dimensional column vector, $\mathbf{p} = [p_1,p_2,...,p_n]^T \in \mathbb{R}^n$ is a $n$-dimensional column vector, $\ln(\cdot)$ is an element-wise natural logarithm function, $\text{sigmoid}(z) = \frac{1}{1 + e^{-z}}$ is an element-wise sigmoid function, and $\circ$ is an element-wise product operator. Use this new information to reformulate the gradients for w and b in matrix form.

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Question 4. Decision Tree

As we recall, the underlying uncertainty of a variable is measured by entropy

\begin{align*}
H(X) = -\sum_iP(X=x_i)\log_2(P(X=x_i))
\end{align*}

When building a tree, one way to decide how to group the values is with the following formula:

\begin{align*}
\arg \max _B G(S, B)=H(S)-\sum_{i=1}^t \frac{\left|S_i\right|}{|S|} H\left(S_i\right)
\end{align*}

where S is the state of the variable of interest, B represents the split that we choose. The larger the value we get the better fit that tree that we made.

You are conducting an fMRI study and discover that when the subject is performing a certain task, some areas in brain are more activated while others are more inhibited. For this experiment, voxels sampled from the region of interest are labeled as either 'excitatory' (green circle) or 'inhibitory' (red triangle). Now you wonder whether we can predict task-induced activation based on certain features of a region, and you hypothesize that feature A and feature B are the most relevant. 

![Q4_DT](imgs/Q4_DT.png)

<!-- BEGIN QUESTION -->

## Question 4.1 
Use the formula to calculate the entropy of activation (i.e. excitation or inhibition) in the region of interest. Show your work. Round your answer to 3 decimal places (0.001).

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 4.2
Now we want to determine whether a region is excitatory or inhibitory based on feature A, i.e. a region is predicted excitatory/inhibitory when its feature A is above certain value. What will be the best split? Give your answer as 'The split should be at feature A = ...(1/2/3/4)', and compute the gain. Round your answer to 3 decimal places (0.001).

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 4.3
Similarly, what will be the best split based on feature B? Give your answer as 'The split should be at feature B = ...(1/2)', and compute the gain. Round your answer to 3 decimal places (0.001).

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 4.4
When constructing a decision tree classifier, what are some good and bad rules of thumb and why?

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# End of A3

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Please make sure to see the output of the gradescope autograder. You are responsible for waiting and ensuring that the autograder is executing normally for your submission. Please create a campuswire post if you see errors in autograder execution.

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['imgs'])
```

Tests for MCQ_LR_1 (points: 0.2)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q1_1, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for MCQ_LR_2 (points: 0.2)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q1_2, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for MCQ_LR_3 (points: 0.2)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q1_3, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for MCQ_LR_4_1 (points: 0.2)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q1_4_1, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for MCQ_LR_4_2 (points: 0.2)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(q1_4_2, all_options), "Is your answer within the option of A/B/C/D?"
```
