```python
# Initialize Otter
import otter
grader = otter.Notebook("A6.ipynb")
```

# Assignment 6


## **Due: May 17th (Wednesday), 2023, 11:59pm (PT)**

### **Instructions:**

Your Jupyter notebook assignment will often have 3 elements: written answers, code answers, and quiz answers. For written answers, you may insert images of your handwritten work in code cells, or write your answers in markdown and LaTeX. For quiz answers, your `record.txt` file will record your answer choices in the quiz modules for submission. Both your quiz answers and code answers will be autograded on Gradescope. This assignment does not have the quiz portion.

For all elements, DO NOT MODIFY THE CELLS. Put your answers **only** in the answer cells given, and **do not delete cells**. If you fail to follow these instructions, you will lose points on your submission.

Make sure to show the steps of your solution for every question to receive credit, not just the final answer. You may search information online but you will need to write code/find solutions to answer the questions yourself. You will submit your .ipynb file and record.txt to gradescope when you are finished.

### **Late Policy:**

Late assignments will be accepted at 75% credit up to one week late. Consult the syllabus for more info on the late policy.

### How to Include Your Math Written Answer?

You could use inline $\LaTeX$ in markdown (recommended) or use markdowns' include image functionality to submit your written responses.

#### $\LaTeX$ (recommended)
[Here is a fantastic tutorial from CalTech about using $\LaTeX$ in Jupyter Notebook.](http://chebe163.caltech.edu/2018w/handouts/intro_to_latex.html). You could also find various $\LaTeX$ tutorials and cheat sheets online.

#### Include Images
If you are still getting familiar with using LaTeX, handwrite the response on paper or the stylus. Take a picture or screenshot of your answer, and include that image in the Jupyter Notebook. Be sure to include that image in the `\imgs` directory. Let's say you have your Q1 response saved as `imgs/Q1.png`; the markdown syntax to include that image is `![Q1](imgs/Q1.png)`. 

## Important Notice

You must check both submission output on the gradescope (`Assignment 6` and `Assignment 6 - Manual Grading`) correctly reflects your work and responses. If you notice inconsistencies between your notebook and the Manual Grading portion, you need to make a campuswire post, and we can help you with that.

**Other**

If you are not feeling comfortable with the programming assignments in this homework, it might help to take a look at [https://github.com/UCSD-COGS108/Tutorials](https://github.com/UCSD-COGS108/Tutorials).

# Question 1: Cross Validation

Cross-validation is implemented in `sklearn`. In this question, you will use `GridSearchCV` from the `model_selection` package to find optimal hyperparameters settings for logistic regression classifiers. Implement the function `cross_validation_logistic_regression` and test it on the breast cancer data. 

_Points:_ 0.5

```python
from sklearn.model_selection import GridSearchCV
from sklearn.linear_model import LogisticRegression

def cross_validation_logistic_regression(X_data, Y_data, k_fold, hyperparams):
    """ test the hyperparams for a logistic regression classifier 
    
    Args:
        X_data (np.ndarray): X data (N x m, where m is the number of features)
        Y_data (np.ndarray): Y data (N x 1) 
        k_fold (int): k for k-fold cross validation
        hyperparams (dict): a dictionary of hyperparameters to test, mapping 
            the name of a hyper-parameters to the list of values to test
    
    Returns:
        best_params (dict): the paramters that gives best performance on the hold-out data
        mean_train_score (np.ndarray): mean of test score
        std_train_score (np.ndarray): standard deviation of test score
        HINTS: they can be read from the attribute `cv_results_` of a `GroundSearchCV` agent. 
    """
    # initialize a logistic regression classifier
    log_reg = ...
    
    # use GridSearchCV to do cross validation
    # for 'scoring' (metrics to measure performance) use accuracy
    
    # get the paramters that gives best performance on the hold-out data
    best_params = ...
    
    # get the mean and standard deviation of test score
    mean_train_score = ...
    std_train_score = ...
    
    return best_params, mean_train_score, std_train_score
```

```python
import numpy as np
import warnings
from sklearn.datasets import load_breast_cancer
from sklearn.exceptions import ConvergenceWarning

# data
cancer_data = load_breast_cancer()
cancer_data_X = cancer_data.data[:200]
cancer_data_Y = cancer_data.target[:200]

# paremeters to tune
params = {
    'C':[0.01, 0.5, 10],
    'solver': ['lbfgs', 'newton-cholesky'],
}

# 5 fold cross-validation
with warnings.catch_warnings():
    warnings.simplefilter("ignore", category=ConvergenceWarning)
    best_params, mean_scores, std_scores = cross_validation_logistic_regression(
        X_data=cancer_data_X,
        Y_data=cancer_data_Y, 
        k_fold=5, 
        hyperparams=params)

print(f'The optimal parameters settings: {best_params}')
print(f'accuracies: {np.round(mean_scores, decimals=3)}')
print(f'(with std: {np.round(std_scores, decimals=3)})')
```

```python
grader.check("cross_validation")
```

# Question 2: SVM

Consider a dataset ${(xi, yi), i = 1, 2, ..., 10}$ where $x = (x_1, x_2)$ and $y \in {−1, +1}$, as shown in the figure below. Suppose we have trained a support vector machine (SVM) classifier on the dataset, which has the decision boundary $l :w^Tx + b = 0$. The SVM classifier is optimized as following:

\begin{align*}
Find &: \arg\min_w = \frac{1}{2} ||w||_2^2 \\
\text{Subject to} &, \forall i : w^Tx_i + b \geq +1, \text{if } y_i = +1, \\
& w^Tx_i + b \leq -1, \text{if } y_i = −1.
\end{align*}

We define $l^+: w^Tx_i + b = +1$ as the positive plane and $l^-: w^Tx_i + b = -1$ as the negative plane. After training the SVM classifier using the gradient descent method for several iterations to reach an intermediate state, we have obtained the positive plane and the negative plane, which are indicated as solid lines in the Figure below.

**Hint**:

$X_1=(4.0, 0.5), X_2(2.0, 2.0)$ are on the negative plane

$X_3=(2.0, 3.0)$ is on the positive plane

![arg_max](imgs/svm_q2.png)

<!-- BEGIN QUESTION -->

2.1 Calculate the w and b from the positive plane and negative plane.

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

2.2 Calculate the size of the margin.

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->


# Question 3: SVM gradient

Given a training dataset $S_{\text{training}} = \{(x_i, y_i), i = 1, . . . , n\}$, we wish to optimize the loss $L(w, b)$ of a linear SVM classifier:

$$L(w, b) = \frac{1}{2}||w||_2^2 + C\sum_{i=1}^{n} 1 − y_i(w^T x_i + b)_{+} $$


where $(z)_+ = max(0, z)$ is called the rectifier function and C is a scalar constant.

The optimal weight vector $w^∗$ and the bias $b^*$ used to build the SVM classifier are defined as follows:

$$ w^∗, b^∗ = \arg \min_{w,b} = L(w, b)$$

In this problem, we attempt to obtain the optimal parameters $w∗$ and $b∗$ by using a standard gradient descent algorithm.

**[Hint]**: To find the derivative of $L(w, b)$, please consider two cases:

(a) $1 − y_i(w^T x_i + b) \geq 0$

(b) $1 − y_i(w^T x_i + b) < 0$

<!-- BEGIN QUESTION -->

## 3.1

find the derivative $\frac{\partial L(w, b)}{\partial w}$

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## 3.2

find the derivative $\frac{\partial L(w, b)}{\partial b}$

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

# Question 4: SVM margin

As shown in the Figure below, given $w$ and $b$ we have the decision boundary (marked as a black line) defined as follows:

$$w^T x + b = 0$$

In parallel to the decision boundary, we have the positive plane (marked as a red line) as

$$w^T x + b = +1$$

and the negative plane (marked as a blue line) defined as:

$$w^T x + b = -1$$

We pick an arbitrary point x on the negative plane, and draw a purple line that (i) passes x and (ii) is perpendicular to the negative plane. The intersection between this purple line and the positive plane can be denoted as $x^+$. Thus, 

we have the following equation that describes the relationship between $x^+$ and $x^-$:

$$x^+ = x^- +\lambda w$$

where $\lambda \in R$ is an undetermined scalar. The margin $M$ is defined as the distance
between the positive and the negative planes, as:

$$M = ||X^+ - X^-||_2 = \sqrt{<\lambda w, \lambda w>}$$

Please derive the following:

$$M = \frac{2}{\sqrt{<w, w>}}$$

![arg_max](imgs/svm_margin.png)

**[Hint]**: First try to represent $\lambda$ in the form of $w$.

_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Q5: SVM concept questions

For multiple choice questions, write your solution as a list of strings by replacing the "..."

Which one below best describes what support vectors are:

A. The training samples that determine the positive and the negative planes.

B. The decision boundary.

C. The test samples that determine the positive and the negative planes.

D. The positive and the negative planes.

_Points:_ 0.3

```python
MCQ_1 = ...
```

```python
grader.check("svm_concepts_1")
```

Consider a dataset $S = \{(x_i y_i), i = 1, 2, ..., n\}$ where $x = [x_1, x_2]^T$ and $y \in \{−1, +1\}$. Here we try to learn a linear SVM classifier parameterized by weight vector $w$ and bias $b$ to fit the dataset $S$. Which one is the appropriate loss function for the linear SVM classifier from the options below:

A. $L(w, b) = \frac{1}{2}w^Tw - C \times \sum_{i=1}^{n} \max(0, 1 − y_i \times (w^Tx_i + b))$

B. $L(w, b) = \frac{2}{\sqrt{w^Tw}} - C \times \sum_{i=1}^{n} \max(0, 1 − y_i \times (w^Tx_i + b))$

C. $L(w, b) = \frac{1}{2}w^Tw + C \times \sum_{i=1}^{n} \max(0, 1 − y_i \times (w^Tx_i + b))$

D. $L(w, b) = \frac{2}{\sqrt{w^Tw}} + C \times \sum_{i=1}^{n} \max(0, 1 − y_i \times (w^Tx_i + b))$

_Points:_ 0.3

```python
MCQ_2 = ...
```

```python
grader.check("svm_concepts_2")
```

In the following figure, We use a black line to mark the decision boundary of a linear SVM classifier parameterized by weight vector w and bias b. If we start to increase the margin of the classifier, what would happen to w?

![svm_margin](imgs/svm_margin.png)

A. w will have smaller magnitude.

B. w will have greater magnitude.

C. The magnitude of w does not change.

D. None of above.

_Points:_ 0.3

```python
MCQ_3 = ...
```

```python
grader.check("svm_concepts_3")
```

# End of A6

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Please make sure to see the output of the gradescope autograder. You are responsible for waiting and ensuring that the autograder is executing normally for your submission. Please create a campuswire post if you see errors in autograder execution.

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['imgs'])
```

Tests for cross_validation (points: 0.5)
```
>>> # sanity check
>>> assert isinstance(best_params, dict), "best_params should be a dictionary of hyper-parameters"
>>> assert 'C' in best_params, "C missing from the hyperparameters dict"
>>> assert 'solver' in best_params, "solver missing from the hyperparameters dict"
>>> assert isinstance(mean_scores, np.ndarray) and mean_scores.shape == (6,), "mean_scores should be an np array of length 6"
>>> assert isinstance(std_scores, np.ndarray) and std_scores.shape == (6,), "std_scores should be an np array of length 6"
```


Tests for svm_concepts_1 (points: 0.3)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(MCQ_1, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for svm_concepts_2 (points: 0.3)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(MCQ_2, all_options), "Is your answer within the option of A/B/C/D?"
```


Tests for svm_concepts_3 (points: 0.3)
```
>>> # sanity check
>>> all_options = ["A", "B", "C", "D"]
>>> check_valid = lambda ans, all_options: all([chosen in all_options for chosen in ans])
>>> assert check_valid(MCQ_3, all_options), "Is your answer within the option of A/B/C/D?"
```
